# Exercise 0.3: Numerical linear algebra methods

---

In this exercise, you will get familiar with `numpy.array` and `numpy.linalg`.

Consider the Heisenberg model of **two** spins,

\begin{equation*}
    \hat{H} = \sum_{\alpha=x,y,z} \hat{\sigma}_1^\alpha \hat{\sigma}_2^\alpha = \hat{\sigma}_1^x \otimes \hat{\sigma}_2^x + \hat{\sigma}_1^y \otimes \hat{\sigma}_2^y + \hat{\sigma}_1^z \otimes \hat{\sigma}_2^z\,,
\end{equation*}

where the Pauli matrices are given as 

\begin{align*}
    \hat{\sigma}^x = \begin{pmatrix}
    0 & 1\\
    1 & 0
    \end{pmatrix}, \qquad
    \hat{\sigma}^y = \begin{pmatrix}
    0 & -i\\
    i & 0
    \end{pmatrix}, \qquad
    \hat{\sigma}^z = \begin{pmatrix}
    1 & 0\\
    0 & -1
    \end{pmatrix}.
\end{align*}

---

```python
### Necessary Python libraries
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
### If you wish, you can import more libraries here
```

---

## Exercise 0.3-a

Construct the Pauli matrices as numpy arrays `sigma_x`, `sigma_y`, and `sigma_z`.

<!-- grade_id: cell-a6f3ec757d672c63, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sigma_x = np.array([[0,1],[1,0]], dtype = complex)
sigma_y = np.array([[0, -1j], [1j, 0]], dtype = complex)
sigma_z = np.array([[1, 0], [0, -1]], dtype = complex)
```

<!-- grade_id: cell-01870e05e3d7b322, points: 0, grade: true, solution: false -->
```python
# 0.2-a tests
assert type(sigma_x) is np.ndarray, "sigma_x must be a numpy array."
assert type(sigma_y) is np.ndarray, "sigma_y must be a numpy array."
assert type(sigma_z) is np.ndarray, "sigma_z must be a numpy array."
assert sigma_x.shape == (2, 2), "sigma_x must be of shape (2, 2)."
assert sigma_y.shape == (2, 2), "sigma_y must be of shape (2, 2)."
assert sigma_z.shape == (2, 2), "sigma_z must be of shape (2, 2)."
```

---

## Exercise 0.3-b

Construct the Hamiltonian of the Heisenberg model as a numpy array `H`.

> **Hint:** For the *kronecker product* $A\otimes B$ you may use `np.kron(A, B)`

<!-- grade_id: cell-4bbd4b0df60a670b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sigma = [sigma_x, sigma_y, sigma_z]

H = np.kron(sigma_x, sigma_x) + np.kron(sigma_y,sigma_y) + np.kron(sigma_z, sigma_z)

```

<!-- grade_id: cell-22e0600e56cfde02, points: 0, grade: true, solution: false -->
```python
# 0.3-b tests
assert type(H) is np.ndarray, "H must be a numpy array."
assert H.shape == (4, 4), "H must be of shape (4, 4)."
# check some of the entries of H
assert np.isclose(H[0, 0], 1).item()
assert np.isclose(H[2, 1], 2).item()
assert np.isclose(H[2, 2], -1).item()
assert np.isclose(H[3, 1], 0).item()
```

---

## Exercise 0.3-c

Diagonalize the Hamiltonian. What are the eigenvalues and eigenstates?

> **Hint:** You can use `numpy.linalg.eigh` to obtain eigenvalues and eigenvectors of hermitean matrices.

<!-- grade_id: cell-f1d1ee46d187f524, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
eigenvalues, eigenvectors = np.linalg.eigh(H)
print(1/np.sqrt(2))
print(eigenvalues)
print(eigenvectors)
```

Output:
```
[-3.  1.  1.  1.]
[[ 0.        +0.j  1.        +0.j  0.        +0.j  0.        +0.j]
 [ 0.70710678+0.j  0.        +0.j  0.70710678+0.j  0.        +0.j]
 [-0.70710678+0.j  0.        +0.j  0.70710678+0.j  0.        +0.j]
 [ 0.        +0.j  0.        +0.j  0.        +0.j  1.        +0.j]]
```

<!-- grade_id: cell-028df1031b4ed387, points: 0, grade: true, solution: false -->
```python
# 0.3-c tests
assert eigenvalues.size == 4
assert eigenvectors.shape == (4, 4)
```

---

## Exercise 0.3-d

Construct the operator for the total magnetization in $z$ direction,

\begin{equation*}
    M = \sigma_1^z \otimes \mathbb{1}_2  +\mathbb{1}_1 \otimes \sigma_2^z \,,
\end{equation*}

and compute its expectation value and variance $\langle M^2\rangle-\langle M \rangle^2$ for for the ground state of $H$ as well as for the state $\frac{1}{\sqrt{2}}(\ket{\uparrow \uparrow}+\ket{\downarrow \downarrow})$.

<!-- grade_id: cell-7cc3b811fcb4c3bc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def exp_val(M, state):
    return np.real(state.conj().T @ M @ state)
def var(M, state):
    return np.real(exp_val(M @ M, state) - (exp_val(M, state))**2)

id = np.eye(2)

M = np.kron(sigma_z, id) + np.kron(id, sigma_z)

print(f"Ground state of Hamiltonian : {eigenvectors[0]}")

#Ground state of H, eigenvector with the lowest eigenvalue
print(f"Expectation value of ground state : {exp_val(M,eigenvectors[0])}")
print(f"Variance of ground state : {var(M, eigenvectors[0])}")

#Given state
state = np.array( [1/np.sqrt(2), 0, 0, 1/np.sqrt(2)] )
print(f"Expectation value of given state : {exp_val(M,state)}")
print(f"Variance of given state : {var(M,state)}")
print('a')

```

Output:
```
Ground state of Hamiltonian : [0.+0.j 1.+0.j 0.+0.j 0.+0.j]
Expectation value of ground state : 0.0
Variance of ground state : 0.0
Expectation value of given state : 0.0
Variance of given state : 3.999999999999999
```

```python

```