# Exercise Sheet 0: Introduction

---

The purpose of this exercise sheet is to give a brief review over some of the basic and necessary functions for this course in python and to introduce our grading scheme.

Please make sure you can solve all the problems below with ease. If you are new to python you may get a very brief (and incomplete) introduction by going through the `intro_python.ipynb` notebook on the gitlab.

> **Note:** <u>This exercise sheet is not counted for the bonus.</u> Therefore, you do not need to show or send the solutions.

---

# Exercise 0.1: Introduction to `nbgrader`

To grade your exercises, we use a grading tool called [**nbgrader**](https://nbgrader.readthedocs.io/en/stable/). The tool supports both automatically and manually graded tasks. Because of the automatic grading some care is required to assure your assignments are graded correctly.

We will upload each assignment as a jupyter notebook on our [gitlab](https://gitlab.lrz.de/cmt-computational-physics/tutorials_2026#). You will work on the assignment by writing code in this notebook. 

To submit your solution, upload your edited notebook on the [moodle page](https://www.moodle.tum.de/course/view.php?id=116136). **Only upload jupyter notebook (`.ipynb`) files. We will not grade any other file types**.

Some weeks assignments will have multiple exercises, each exercise given as one jupyter notebook. In this case, you should also **upload each notebook seperately** (not in a `.zip` file).

**Do not change the filenames of the notebooks.**

Each assignment notebook consists of *cells*. There are different cell types:

* *Read-only* cells with the exercise description
* Editable code cells with the marker `# YOUR CODE HERE`. Write your solutions (as code) in these cells.
* *Read-only* cells with automatic tests. You can execute these cells to check if your solution passes the tests. Usually we will also have some additional *hidden tests* that are not visible to you, but will be executed after you hand in the assignment.
* Editable *markdown* cells with the marker `YOUR ANSWER HERE`. These cells are for written answers. You can use standard latex to write math symbols, e.g. `$f(x)=x^2$` will produce $f(x)=x^2$.

Depending on which program you use to work on the assignment (e.g. jupyter lab, VSCode, etc.), you may be able to edit the *read-only* cells. Keep in mind that the grading tool will replace these edited cells with the original cells after you submit your solution. 

We recommend working on the assignments using jupyter lab, because it correctly identifies the *read-only* cells and prevents you from editing them. To use jupyter lab, open the console, move into the correct folder and execute the command `jupyter lab`, which will open jupyter lab in your browser.

You are allowed to add additional cells to the notebook as required. To make grading easier we ask you to add additional cells only above the `# YOUR CODE HERE` cells, if possible.

For the grading we enforce a **time limit** on your solutions. **Each individual cell must run in less than $5$ minutes and the full notebook must run in less than $15$ minutes** on a standard work station in the CIP pool.

Below is an example of how the assignment notebooks will look like.

```python
### Necessary Python libraries
import numpy as np
import matplotlib.pyplot as plt
```

```python
### If you wish, you can import more libraries here
```

---

## Exercise 0.1-a

Implement the function `f(x)` that computed $f(x) = x^2$. 

<!-- grade_id: cell-04b06278d631b983, points: 0, grade: false, solution: true -->
```python
def f(x):
    # YOUR CODE HERE
```

<!-- grade_id: cell-7c20b8a282716488, points: 0, grade: true, solution: false -->
```python
# 0.1-a tests
assert f(2) == 4, "f(x) is not implemented correctly"
assert f(-2) == 4, "f(x) is not implemented correctly"
```

---

## Exercise 0.1-b

What is the derivative $\frac{\text{d}}{\text{d}x}f(x)$? Write your answer in the markdown cell below.

<!-- grade_id: cell-b3954ccf143e2a83, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

---

## Exercise 0.1-c

Plot $f(x)$ and $\frac{\text{d}}{\text{d}x}f(x)$ in the interval $x\in[-2, 2]$.

<!-- grade_id: cell-ad226c3636bc76d6, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
```