# Exercise 0.2: Plotting

In this exercise, you will get familiar with `numpy.array` and the `matplotlib` package.

```python
### Necessary Python libraries
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
### If you wish, you can import more libraries here
```

---

We start by plotting the function $f(x)=x^2$ for $x\in[-1, 1]$.

## Exercise 0.2-a

Create a numpy array `X` of evenly distributed points in $[-1, 1]$. For example, $X=[-1., -0.9, -0.8,\ldots,1.]$

<!-- grade_id: cell-0cb3d001c32e1633, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-6931b0a35cc31600, points: 1, grade: true, solution: false -->
```python
# 0.2-a tests
assert type(X) is np.ndarray, "X must be a numpy array."
assert X.ndim == 1, "X must be a one-dimensional array."
assert np.isclose(X[0], -1).item(), f"The first element of X must be -1, but is {X[0]}"
assert np.isclose(X[-1], 1).item(), f"The last element of X must be 1, but is {X[-1]}"
```

---

## Exercise 0.2-b

Create a numpy array $Y$ of $f(x_i)$ with point $x_i$ from the array $X$ we created.


<!-- grade_id: cell-a637ad8983a6a8ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-0b5a12bdf1e522cb, points: 1, grade: true, solution: false -->
```python
# 0.2-b tests
assert type(Y) is np.ndarray, "Y must be a numpy array."
assert Y.shape == X.shape, "Y must have the same shape as X."
assert np.isclose(Y[0], 1).item(), f"The first element of Y must be 1, but is {Y[0]}"
assert np.isclose(Y[-1], 1).item(), f"The last element of Y must be 1, but is {Y[-1]}"
```

---

## Exercise 0.2-c

Plot the function.

> **Hint:** Use `matplotlib.pyplot.plt`.

<!-- grade_id: cell-3c420f9f3cc0205c, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

---

## Exercise 0.2-d

Repeat step (a)-(c), now using the function $f(x) = e^{-x}$ for $x\in[0, 10]$.

<!-- grade_id: cell-c88289e9a8800ab8, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

---

## Exercise 0.2-e

Plot the same function again, using now a logarithmic scale for the $y$-axis.

>**Hint:** You may use `matplotlib.pyplot.semilogy`.

<!-- grade_id: cell-ccf4549c06895b63, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

---

Suppose in additional to the data $Y$, we also have the standard deviation $\Delta Y$ to the data resulting from some random error.

---

## Exercise 0.2-f

Create a numpy array `delta_Y` and put in some small random numbers. 

> **Note**: In this course, we will use the random sampling functionality using the `numpy.random.Generator` class, [which was introduced in numpy version `1.17.0`](https://numpy.org/doc/stable/reference/random/new-or-different.html) and replaced the legacy `numpy.random.RandomState`.
>
> To generate random numbers, first create an instance of `numpy.random.Generator`, e.g. using `rng = numpy.random.default_rng(seed)`, with `seed` an arbitrary integer. You can then draw random numbers from this generator using e.g. `rng.random()` (uniform distribution $\in[0,1)$) or `rng.normal()` (normal distribution).

<!-- grade_id: cell-c8dccf987049b5b0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-fde95fb2d129bfe4, points: 1, grade: true, solution: false -->
```python
# 0.2-f tests
assert type(delta_Y) is np.ndarray, "delta_Y must be a numpy array."
assert delta_Y.shape == X.shape, "delta_Y must have the same shape as X."
assert np.all(delta_Y >= 0), "All errors must be larger or equal to 0."
```

<!-- grade_id: cell-a8e8d076752af7bb, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```