# Exercise 2: Computational Complexity of Matrix Operations (*total: 15 points*)

---

**Note**: Because the 1st of May is a holiday, there will be no tutorial on that day. This exercise will however still count for the bonus. The submission deadline is the 4th of May at 1pm.

---

```python
### Necessary Python libraries
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
import numpy as np
rng = np.random.default_rng(seed=3141592653)
import time
from matplotlib import pyplot as plt
%matplotlib inline
```

```python
### If you wish, you can import more libraries here
```

---

## Exercise 2-a (2pt)

Take a look at the code below and undestand how the `timeit(..)` function works. We will use this function to evaluate the runtime of matrix-vector and matrix-matrix multiplication with different implementations. Complete the functions `matrix_vector_naive(mat, vec)` and `matrix_matrix_naive(mat1, mat2)` to implement matrix-vector and matrix-matrix multiplication using only standard python `for` loops.

```python
def timeit(func_setup, func_eval, ns=[50, 100, 200, 400, 600, 1000, 2000, 4000, 6000, 10000], max_time=1.):
    print("starting run")
    ts = []
    for n in ns:
        print("  timing n={0}".format(n))
        args = func_setup(n)
        t0 = time.time()
        res = func_eval(*args)
        t1 = time.time()
        ts.append(t1-t0)
        if t1 - t0 > max_time:
            print(f"skip remaining since it took > {max_time:f}s")
            break
    return ns[:len(ts)], ts
```

```python
def setup_matrix_vector(n):
    mat = rng.normal(size=(n,n))
    vec = rng.normal(size=n)
    return mat, vec

def setup_matrix_matrix(n):
    mat1 = rng.normal(size=(n,n))
    mat2 = rng.normal(size=(n,n))
    return mat1, mat2
```

<!-- grade_id: cell-4c47854fc10e2ad6, points: 0, grade: false, solution: true -->
```python
def matrix_vector_naive(mat, vec):
    # YOUR CODE HERE

def matrix_matrix_naive(mat1, mat2):
    # YOUR CODE HERE
```

<!-- grade_id: cell-378ae7849f3c6514, points: 2, grade: true, solution: false -->
```python
# 2-a tests
# Test matrix_vector_naive
rng = np.random.default_rng(123)
mat_test, vec_test = setup_matrix_vector(3)
result_test = matrix_vector_naive(mat_test, vec_test)
assert np.isclose(result_test, np.array([-1.6821392, -0.8537391, 0.74095162])).all(), "matrix_vector_naive() is not implemented correctly!"
# Test matrix_matrix_naive
mat1_test, mat2_test = setup_matrix_matrix(3)
result_test = matrix_matrix_naive(mat1_test, mat2_test)
assert np.isclose(result_test, np.array([
    [1.55482175, 2.32857038, 0.48818521],
    [0.56036788, -0.27260392, 2.1141489],
    [-2.37409532, -1.39667842, 0.07036474]])).all(), "matrix_matrix_naive() is not implemented correctly!"
del mat_test, vec_test, mat1_test, mat2_test, result_test
```

---

## Exercise 2-b (2pt)

Complete the functions `matrix_vector_numpy(mat, vec)` and `matrix_matrix_numpy(mat1, mat2)` to implement the same operations but this time using `numpy`

<!-- grade_id: cell-f8e3cacfd7edbe44, points: 0, grade: false, solution: true -->
```python
def matrix_vector_numpy(mat, vec):
    # YOUR CODE HERE

def matrix_matrix_numpy(mat1, mat2):
    # YOUR CODE HERE
```

<!-- grade_id: cell-9d93676c9aeb7cc2, points: 2, grade: true, solution: false -->
```python
# 2-b tests
rng = np.random.default_rng(123)
mat_test, vec_test = setup_matrix_vector(3)
result_test = matrix_vector_numpy(mat_test, vec_test)
assert np.isclose(result_test, np.array([-1.6821392, -0.8537391, 0.74095162])).all(), "matrix_vector_numpy() is not implemented correctly!"
# Test matrix_matrix_naive
mat1_test, mat2_test = setup_matrix_matrix(3)
result_test = matrix_matrix_numpy(mat1_test, mat2_test)
assert np.isclose(result_test, np.array([
    [1.55482175, 2.32857038, 0.48818521],
    [0.56036788, -0.27260392, 2.1141489],
    [-2.37409532, -1.39667842, 0.07036474]])).all(), "matrix_matrix_numpy() is not implemented correctly!"
del mat_test, vec_test, mat1_test, mat2_test, result_test
```

---

## Exercise 2-c (0pt)

Evaluate the runtime of all implemented funtions using the `timeit(...)` function.

<!-- grade_id: cell-6b6c41faab54718b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

---

## Exercise 2-d (4pt)

How many floating point operations (additions and multiplications) need to be performed for a matrix-vector and matrix-matrix multiplication of dimension $n$ in the naive implementation? How much longer does a matrix-matrix-multiplcation of size $2n$ take compared to size $n$?

> **Hint:** We are only interested in the leading order in $n$.


<!-- grade_id: cell-d883ac551bb72fcf, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 2-e (3pt)

Plot the runtimes and analytical prediction for the number of floating point operations for matrix-vector and matrix-matrix multiplication, for both the naive and the numpy implementation, in log-log scale. Do you observe what you expected? For what prefactor does the analytical prediction match the runtimes obtained using the naive approach? What is the meaning of this prefactor? Write your answers in the markdown cell below.


<!-- grade_id: cell-b785eaf7fa0a2584, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-7cb51d68038d0ad4, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---

We now consider vectors and matrices with complex entries.

## Exercise 2-f (2pt)

How much longer do you expect each multiplication to take with complex arguments instead of real ones? How much longer for each addition? Use you results to compute how much longer a matrix-vector and a matrix-matrix multiplication  takes with complex arguments instead of real ones. Write your answers in the markdown cell below.


<!-- grade_id: cell-67e81fe24ad1e68b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

---

## Exercise 2-g (1pt)

Plot the runtimes comparing real and complex arguments as in part d) and compute the average ratio $t_\mathrm{complex}/t_\mathrm{float}$ (in the region where it makes sense). Does your result agree with your analytical prediction?

> Hint: You can time the functions using the numpy implementation from part b), using the setup functions `setup_matrix_vector_complex` and `setup_matrix_matrix_complex`provided below.


```python
def setup_matrix_vector_complex(n):
    return rng.normal(size=(n,n)) + 1j*rng.normal(size=(n,n)), rng.normal(size=n) + 1j*rng.normal(size=n)

def setup_matrix_matrix_complex(n):
    return rng.normal(size=(n,n)) + 1j*rng.normal(size=(n,n)), rng.normal(size=(n,n)) + 1j*rng.normal(size=(n,n))
```

<!-- grade_id: cell-935fe74f3330d407, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```