# Exercise 1.2: Importance sampling with Monte Carlo (*total: 8 points*)

---

The goal of this exercise is to find a Monte Carlo estimate of the integral

\begin{equation*}
    I = \int_{a=0}^{\infty} \mathrm{d} x \frac{\mathrm{e}^{-x}}{1+(x-1)^2}
\end{equation*}

(It is fine to cut off the upper limit at some value e.g. $b=10$).

---

```python
### Necessary Python libraries
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
import numpy as np
import matplotlib.pyplot as plt
import scipy.integrate
%matplotlib inline
```

```python
### If you wish, you can import more libraries here
```

---

## Exercise 1.2-a (1pt)

Implement the function `f(x)` which simply computes the value of
\begin{equation*}
    f(x) = \frac{\mathrm{e}^{-x}}{1+(x-1)^2}.
\end{equation*}

Make sure that your implementation works when a 1D `numpy` array containing $N$ values $x$ at once is passed as an argument.

<!-- grade_id: cell-5485fdf71e9b094a, points: 0, grade: false, solution: true -->
```python
def f(x):
    # YOUR CODE HERE
    if np.isscalar(x):
        one_vector = 1
    else:
        one_vector = np.ones(len(x))
    return np.exp(-x) / (one_vector + (x - one_vector) ** 2)
```

<!-- grade_id: cell-612ce454ab9c0c35, points: 1, grade: true, solution: false -->
```python
# 1.2-a tests
result_test = f(0.67)
assert type(result_test) is float or type(result_test) is np.float64, f"Expected a result of type float, but the result type is {type(result_test)}"
assert np.isclose(result_test, 0.4614560174826787).item(), f"Your implementation of f(x) does not return the correct value for x=0.67."
result_test = f(np.array([1., 2., 3.]))
assert type(result_test) is np.ndarray, f"Expected a result of type np.ndarray, but the result type is {type(result_test)}"
assert np.isclose(result_test, [0.36787944, 0.06766764, 0.00995741]).all(), f"Your implementation of f(x) does not return the correct value for x=[1, 2, 3]."
del result_test
```

---

## Exercise 1.2-b (2pt)

Implement the function `estimate_I_monte_carlo(rng, N, a, b=10)` that calculates $I$ by using that $I = \overline{f} \times (b-a) $ with $f(x) = \frac{\mathrm{e}^{-x}}{1+(x-1)^2} $ (You can find an estimate of $\overline{f}$ by sampling some $x_i \in [a,b)$ from an uniform distribution, and averaging over $f(x_i)$). Get an idea of the error of this estimate by plotting the mean error against the number of samples.

> **Note:** To get a value of $I$ to compare your estimate to you may use `scipy.integrate.quad`.


<!-- grade_id: cell-8f8ce9c852175dcd, points: 0, grade: false, solution: true -->
```python
def estimate_I_monte_carlo(rng, N, a, b=10):
    # YOUR CODE HERE
```

<!-- grade_id: cell-5652fa082feb9b90, points: 1, grade: true, solution: false -->
```python
# 1.2-b tests
result_test = estimate_I_monte_carlo(np.random.default_rng(42), 1000, 0, 10)
assert np.isclose(result_test, 0.7, atol=2.e-1), f"The estimation for N=1000 samples should already be closer to the exact result."
del result_test
```

<!-- grade_id: cell-c14ca82b55ecc898, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

---

In the remainder of the exercise, we will implement importance sampling to improve our estimation. We will follow the steps outlined in the lecture to sample from the probability distribution function (p.d.f)

\begin{equation*}
    g(x) = \alpha \mathrm{e}^{-\alpha x}
\end{equation*}

with $\alpha \approx 1.46$ instead of the uniform distribution over $[a, b]$.

---

## Exercise 1.2-c (2pt)

 Start by implementing `g(x, alpha=1.46)` and plotting $f(x)$ and $g(x)$ over the interval $[0, 10]$ in the same plot. What do you notice? Write your answer in the markdown cell below.


<!-- grade_id: cell-7877e286891cef77, points: 0, grade: false, solution: true -->
```python
def g(x, alpha=1.46):
    # YOUR CODE HERE
```

<!-- grade_id: cell-b888496572ad5bd8, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-40d7f8a05ee426f0, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---

## Exercise 1.2-d (2pt)

Derive the cumulative distribution function (c.d.f) 

\begin{equation*}
    G(x) = \int_0^x \text{d}x^\prime g(x^\prime)
\end{equation*}

and its inverse $G^{-1}(u)$. Implement the functions `inverse_CDF(u, alpha=1.46)`, which computes the value of the inverse c.d.f. and `sample_from_pdf(rng, N, alpha=1.46)`, which produces $N$ samples sampled from the p.d.f. Using the plotting code provided below, produce a histogram of $N\approx 1000$ sampled points to check if your implementation is correct. 


<!-- grade_id: cell-6dcbd0a2dead6a00, points: 0, grade: false, solution: true -->
```python
def inverse_cdf(u, alpha=1.46):
    # YOUR CODE HERE
```

<!-- grade_id: cell-d6f9206e18b68e12, points: 0, grade: false, solution: true -->
```python
def sample_from_pdf(rng, N, alpha=1.46):
    # YOUR CODE HERE
```

<!-- grade_id: cell-b224d24a835ba579, points: 1, grade: true, solution: false -->
```python
# 1.2-d tests
result_test = inverse_cdf(0.2)
assert type(result_test) is float or type(result_test) is np.float64, f"Expected a result of type float, but the result type is {type(result_test)}"
assert np.isclose(result_test, 0.15283804884534913).item()
assert np.isclose(inverse_cdf(0.8), 1.1023547345439044).item()
result_test = sample_from_pdf(np.random.default_rng(42), 100)
assert (type(result_test) is np.ndarray and result_test.shape == (100,) and result_test.dtype == np.float64) or (type(result_test) is list and len(result_test) == 100 and type(result_test[0]) is float), f"result is not a 1D np.ndarray or list of length 100."
del result_test
```

```python
def plot_histogram(points, N_bins=100):
    plt.hist(points, N_bins, label="random sample points", density=True)
    xs = np.linspace(np.min(points), np.max(points), 100)
    plt.plot(xs, g(xs), "--", label="g(x)")
    plt.xlabel(r"$x$")
    plt.legend()
    plt.show()
```

<!-- grade_id: cell-3acecf0d4a2c1b88, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

---

## Exercise 1.2-d (1pt)

Implement the function `estimate_I_monte_carlo_importance_sampling(rng, N, alpha=1.46)`. Compare the estimation with and without importance sampling by recreating the plot from 1.2-a, now additionally including the sampling with importance sampling.


<!-- grade_id: cell-37e42c014b743a36, points: 0, grade: false, solution: true -->
```python
def estimate_I_monte_carlo_importance_sampling(rng, N, alpha=1.46):
    # YOUR CODE HERE
```

<!-- grade_id: cell-707d20adfc2b3256, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```