# Exercise 1.1: Estimate $\pi$ with Monte Carlo (*total: 7 points*)

```python
### Necessary Python libraries
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
import numpy as np
import matplotlib.pyplot as plt
import scipy.special
%matplotlib inline
```

```python
### If you wish, you can import more libraries here
```

---

## Exercise 1.1-a (1pt)

Implement the function `estimate_pi(ng, N)`, which estimates $\pi$ by “shooting” (i.e., drawing random numbers) $N$ times uniformly on a square and counting the number of points hitting a disc target: the ratio of hits to $N$ corresponds to the ratio of the areas of the target to the area you shoot on.

<div align="center">
  <img src="compute_pi_with_MC.svg" alt="drawing" width="200"/>
</div>

<!-- grade_id: cell-4d9c2abdab5d5fa0, points: 0, grade: false, solution: true -->
```python
def estimate_pi(rng, N):
    # YOUR CODE HERE
```

<!-- grade_id: cell-1b786a817007f06d, points: 0, grade: false, solution: true -->
```python
# Testing
# YOUR CODE HERE
```

<!-- grade_id: cell-8abd3d024e3595b6, points: 1, grade: true, solution: false -->
```python
# 1.1-a tests
result_test = estimate_pi(np.random.default_rng(42), 1)
assert type(result_test) is float or type(result_test) is np.float64, f"Expected a result of type float, but the result type is {type(result_test)}"
result_test = estimate_pi(np.random.default_rng(42), 1000)
assert np.isclose(result_test, np.pi, atol=2.e-1), f"The estimation for N=1000 samples should already be close to pi, but is {result_test}."
del result_test
```

---

## Exercise 1.1-b (1pt)

Implement the function `estimate_average_error(rng, N, N_bins)`, estimating the average error for given $N$ by repeating 1.1-a a few ($N_\text{bins}$) times.

> **Optional:** Additionally compute the *variance*
> \begin{equation*}
\text{VAR}(x_1, x_2, \dots, x_{N_\text{bins}}) = \frac{\sum_{n=1}^{N_\text{bins}}|x_n-\bar{x}|}{N_\text{bins}},
\end{equation*}
> where $x_n$ is the estimated value of $\pi$ from bin $n$ and $\bar{x}$ is the average of all bins.

<!-- grade_id: cell-67968540ef55ea4f, points: 0, grade: false, solution: true -->
```python
def estimate_average_error(rng, N, N_bins):
    # YOUR CODE HERE
```

<!-- grade_id: cell-cd66f609e3fd52dc, points: 0, grade: false, solution: true -->
```python
# Testing
# YOUR CODE HERE
```

<!-- grade_id: cell-829858d307970417, points: 1, grade: true, solution: false -->
```python
# 1.1-b tests
result_test = estimate_average_error(np.random.default_rng(42), 1000, 100)
assert type(result_test) is float or type(result_test) is np.float64, f"Expected a result of type float, but the result type is {type(result_test)}"
assert result_test < 0.1, f"1000 samples should result in a mean error < 0.05, but the mean error is {result_test}."
del result_test
```

---

## Exercise 1.1-c (2pt)

Plot the average error versus $N$ on a log-log scale. How does the error scale?


<!-- grade_id: cell-9d96e3fc933a6223, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-f3c7377c1d7de891, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---

Now we want to generalize our problem to $d$-dimensions.

## Exercise 1.1-d (1pt)

Generalize your code to estimate the volume of a $d$-dimensional sphere of radius $r=1$ using Monte Carlo sampling by implementing the function `estimate_volume_d_sphere(rng, N, d)`.

<!-- grade_id: cell-3378d9c248a688ab, points: 0, grade: false, solution: true -->
```python
def estimate_volume_d_sphere(rng, N, d):
    # YOUR CODE HERE
```

<!-- grade_id: cell-951e190664273e2a, points: 0, grade: false, solution: true -->
```python
# Testing
# YOUR CODE HERE
```

<!-- grade_id: cell-4f3204bcb667f20f, points: 1, grade: true, solution: false -->
```python
# 1.1-d tests
result_test = estimate_volume_d_sphere(np.random.default_rng(42), 10000, 3)
assert type(result_test) is float or type(result_test) is np.float64, f"Expected a result of type float, but the result type is {type(result_test)}"
assert np.isclose(result_test, 4.189, atol=0.5)
assert np.isclose(estimate_volume_d_sphere(np.random.default_rng(42), 10000, 4), 4.935, atol=1.0)
del result_test
```

---

## Exercise 1.1-e (2pt)

Plot the average error versus $N$ on a log-log-scale for different values of $d$. How does the computational cost scale with the of dimension $d$ for a fixed value of $N$? How does the error of the estimates scale for $d$ and $N$? Write down your answers in the markdown cell below.

> **Note:** To compute the error of your estimate you can compare to the [exact value](https://en.wikipedia.org/wiki/Volume_of_an_n-ball)
> \begin{equation*}
V_{\text{d-sphere}} = \frac{\pi^{d/2}}{\Gamma\left(\frac{d}{2}+1\right)}.
\end{equation*}
> The value of the $\Gamma$-function can be obtained using `scipy.special.gamma`.

> **Hint:** To see the scaling of the computational cost and the scaling of the error with $d$ at fixed $N$ it may be helpful (but not required) to produce additional plots.


```python
def compute_volume_sphere_exact(d):
    return np.pi**(d/2)/scipy.special.gamma(d/2+1)
```

<!-- grade_id: cell-ad1914b4464e3d43, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-1bd79711edb22006, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE